# Aufgabensatz 1
Zum Start gibt es nun ein paar kleinere Übungen um mit Python warm zu werden

## 1a)
Erstelle eine Funktion, welche jedes zweite Element in einer Liste entfernt und die neue Liste dann zurückgibt

<!-- grade_id: cell-remove_second, points: 0, grade: false, solution: true -->
```python
def remove_every_second_element(in_list: list) -> list:
    r"""
        Die Funktion erhält eine Liste *in_list* und gibt eine Liste zurück, in der jedes zweite Element fehlt.
        
        Beispiel
        ========
        >>> remove_every_second_element([0, 1, 2, 3, 4, 5])
        [0, 2, 4]
        >>> remove_every_second_element(['abc', 'Welle', 'kalt', 'Pferd', 'Apfel'])
        ['abc', 'kalt', 'Apfel']
    """
    # YOUR CODE HERE
    return in_list[::2]
```

Nun ein paar Tests um die Funktionsfähigkeit zu überprüfen. Sollte es zu einem Fehler kommen, wird ein `AssertionError` zurückgegeben. Die angezeigten Tests sind nicht unbedingt alle die ausprobiert werden. Überlegen Sie sich daher selber Tests.

<!-- grade_id: cell-test_remove_second, points: 1, grade: true, solution: false -->
```python
assert remove_every_second_element([0, 1, 2, 3, 4, 5]) == [0, 2, 4]
assert remove_every_second_element(['abc', 'Welle', 'kalt', 'Pferd', 'Apfel']) == ['abc', 'kalt', 'Apfel']
assert remove_every_second_element(list(range(0,10))) == list(range(0, 10, 2))
```

## 1b)
Erstelle eine Funktion, welche drei Strings bekommt und zu einem einzigen String zusammenfügt.
Dabei gilt folgende Regel:
* Aus jedem String wird der erste Buchstabe genommen zusammengeführt und anschließend der zweite aus jedem usw.
* Ist ein String kürzer als die anderen gilt die obige Regel nur noch für zwei, dann einen der Strings

<!-- grade_id: cell-join_strings, points: 0, grade: false, solution: true -->
```python
def join_strings(string_1: str, string_2: str, string_3: str) -> str:
    r"""
        Diese Funktion erhält drei *Strings* und gibt einen *String* zurück. Dabei werden nacheinander die
        Buchstaben am Anfang zum neuen String hinzugefügt. Danach geht es mit dem nächsten Buchstaben weiter.
        Sind ein oder zwei Strings kürzer, wird die Regel auf die Verbliebenen Strings angewendet.
        
        Beispiel
        ========
        >>> join_strings('abc', 'abc', 'abc')
        'aaabbbccc'
        >>> join_strings('Hl l', 'aoWt', 'l,e!')
        'Hallo, Welt!'
    """
    # YOUR CODE HERE
    import itertools
    msg = ""
    for val1, val2, val3 in itertools.zip_longest(string_1, string_2, string_3):
        msg += val1 + val2 + val3
    return msg
    
```

<!-- grade_id: cell-test_join_strings, points: 1, grade: true, solution: false -->
```python
assert join_strings('abc', 'abc', 'abc') == 'aaabbbccc'
assert join_strings('Hl l', 'aoWt', 'l,e!') == 'Hallo, Welt!'
```

## 1c)
Erstelle eine Funktion, welche den Mittelwert, den Median und die Standardabweichung einer Liste von Zahlen als **Tuple** zurückgibt. Nutze dafür, wenn möglich, nur die Python Standardbibliothek.
### Zur Erinnerung
* Mittelwert https://de.wikipedia.org/wiki/Mittelwert
* Median https://de.wikipedia.org/wiki/Median
* Standardabweichung - https://de.wikipedia.org/wiki/Varianz_(Stochastik)

<!-- grade_id: cell-statistics, points: 0, grade: false, solution: true -->
```python
def get_statistics(array: list) -> tuple:
    r"""
        Diese Funktion berechnet den Mittelwert, den Median und die Standardabweichung einer Liste von Zahlen.
        
        Beispiel
        ========
        >>> get_statistics([1, 2, 3, 4, 5, 6, 7, 8, 9])
        (5, 5, 2.581989)
        >>> get_statistics([10, 5, 8, 2, 1, 3, 9])
        [5.42857, 5, 3.33197]
    """
    # YOUR CODE HERE
    lenArray = len(array)
    mean = sum(array) / lenArray
    
    array.sort()
    if (lenArray % 2 == 0): #even
        median = (array[int(lenArray / 2) - 1] + array[int(lenArray / 2)]) / 2
    else: # not even
        median = array[int(lenArray / 2)]
    
    var = 0
    for value in array:
        var += (value - mean) ** 2
    std = (var / (lenArray)) ** 0.5
        
    return mean, median, std
```

<!-- grade_id: cell-test_statistics, points: 1, grade: true, solution: false -->
```python
mean, median, std = get_statistics([1, 2, 3, 4, 5, 6, 7, 8, 9])
assert round(mean, 2) == 5.0
assert round(median, 2) == 5.0
assert round(std, 3) == 2.582
mean, median, std = get_statistics([10, 5, 8, 2, 1, 3, 9])
assert round(mean, 2) == 5.43
assert round(median, 2) == 5.0
assert round(std, 3) == 3.332
```

## 1d)
Erstelle eine Funktion, welche aus einem beliebigen String einen bestimmten Buchstaben komplett entfernt.

<!-- grade_id: cell-letter_filter, points: 0, grade: false, solution: true -->
```python
def filter_letter(string: str, taboo: str) -> str:
    r"""
        Diese Funktion filtert aus einen gegebenen String jedes Auftreten eines Buchstaben heraus. Buchstaben
        können auch Zahlen oder sonstige ASCII Zeichen sein.
        
        Beispiel
        ========
        >>> filter_letter('Hxalxlxox, Wxxexlxt!', 'x')
        'Hallo, Welt!'
    """
    # YOUR CODE HERE
    return string.replace(taboo, "")
```

<!-- grade_id: cell-test_letter_filter, points: 1, grade: true, solution: false -->
```python
assert filter_letter('Hxalxlxox, Wxxexlxt!', 'x') == 'Hallo, Welt!'
assert filter_letter('asdflkbnlkernsdofinlniasdlfn', 'a') == 'sdflkbnlkernsdofinlnisdlfn'
```

## 1e)
Erstelle eine Funktion, welche eine Datein einließt und deren Inhalt übergibt.
### Für die Tests:
Es sollten sich folgende Datein im Verzeichnis dieses Notebooks befinden
* `HalloWelt.txt`
* `Hexeneinmaleins.txt`
* `SPAM.txt`

<!-- grade_id: cell-data_reading, points: 0, grade: false, solution: true -->
```python
def data_reader(path: str) -> str:
    r"""
        Diese Funktion liest den Inhalt der unter *path* hinterlegten Datei und gibt dieses zurück.
        
        Beispiel
        ========
        >>> data_reader('HalloWelt.txt')
        'Hallo, Welt!'
    """
    # YOUR CODE HERE
    s = open(path, "r", encoding='utf-8').read()
    return s
```

<!-- grade_id: cell-test_data_reading, points: 1, grade: true, solution: false -->
```python
assert data_reader('HalloWelt.txt') == 'Hallo, Welt!'
assert data_reader('Hexeneinmaleins.txt') == """Du mußt verstehn!\nAus Eins mach’ Zehn,\nUnd Zwei laß gehn,
Und Drei mach’ gleich,\nSo bist Du reich.\nVerlier’ die Vier!\nAus Fünf und Sechs,\nSo sagt die Hex’,
Mach’ Sieben und Acht,\nSo ist’s vollbracht:\nUnd Neun ist Eins,\nUnd Zehn ist keins.
Das ist das Hexen-Einmal-Eins!"""
assert data_reader('SPAM.txt') == """Spam, Spam, Spam, lovely Spam\nWonderful Spam, Lovely Spam.
Spam, Spam, Spam, magnificent Spam,\nSuperlative Spam.\nSpam, Spam, Spam, wonderous Spam,
Surgical Spam, splendiferous Spam.\nSpam, Spam, Spam, Spaaam!"""
```